// I AM NOT DONE

/*
Difficulty: ⭐⭐⭐⭐
Topic: Enums Quiz

Description:
Let's build a `Shape` enum!
1. Variants: `Circle(f64)`, `Rectangle(f64, f64)`, `Triangle(f64, f64, f64)`.
2. Method `perimeter` which returns the perimeter as `f64`.

Your task is to implement the enum and the method.
*/

// TODO: Define enum Shape

// TODO: Implement perimeter method

fn main() {
    // Uncomment these lines when ready
    /*
    let circle = Shape::Circle(10.0);
    let rect = Shape::Rectangle(10.0, 20.0);
    let tri = Shape::Triangle(10.0, 10.0, 10.0);

    println!("Circle perimeter: {}", circle.perimeter());
    println!("Rectangle perimeter: {}", rect.perimeter());
    println!("Triangle perimeter: {}", tri.perimeter());
    */
}

#[cfg(test)]
mod tests {
    #[test]
    fn test_main_runs() {
        super::main();
    }
}
