// I AM NOT DONE

/*
Difficulty: ⭐⭐⭐
Topic: Enums

Description:
⭐⭐⭐ - Safe unwrapping (unwrap_or, unwrap_or_else)
*/

fn main() {
    // TODO: Fix this code
    let x = "change me";
    println!("Exercise: {}", x);
}

#[cfg(test)]
mod tests {
    #[test]
    fn test_main_runs() {
        super::main();
    }
}
