// I AM NOT DONE

/*
Difficulty: ⭐
Topic: Result - Definition

Description:
`Result` is an enum for returning either success (`Ok`) or failure (`Err`).
It has two variants: `Ok(T)` and `Err(E)`.

Your task is to:
1. Create a variable `success` which is a `Result<i32, String>` containing `Ok(10)`.
2. Create a variable `failure` which is a `Result<i32, String>` containing `Err(String::from("Error"))`.
*/

fn main() {
    // TODO: Create success
    // let success: Result<i32, String> = ...;

    // TODO: Create failure
    // let failure: Result<i32, String> = ...;

    // println!("{:?}, {:?}", success, failure);
}

#[cfg(test)]
mod tests {
    #[test]
    fn test_main_runs() {
        super::main();
    }
}
