// I AM NOT DONE

/*
Difficulty: ⭐
Topic: Ownership

Description:
⭐ - Understand why a variable is invalid after move
*/

fn main() {
    // TODO: Fix this code
    let x = "change me";
    println!("Exercise: {}", x);
}

#[cfg(test)]
mod tests {
    #[test]
    fn test_main_runs() {
        super::main();
    }
}
