// I AM NOT DONE

/*
Difficulty: ⭐⭐⭐⭐⭐
Topic: Lifetimes Quiz

Description:
Implement a struct `Book` that holds a reference to a `title` (string slice) and an `author` (string slice).
Implement a method `get_title` that returns the title.
Ensure all lifetimes are correct.
*/

// TODO: Define struct Book with lifetimes

// TODO: Implement methods

fn main() {
    // let title = String::from("Rust Book");
    // let author = String::from("Steve");

    // let book = Book { title: &title, author: &author };

    // println!("Title: {}", book.get_title());
}

#[cfg(test)]
mod tests {
    #[test]
    fn test_main_runs() {
        super::main();
    }
}
