// I AM NOT DONE

/*
Difficulty: ⭐⭐⭐
Topic: Structs - Unit-Like Structs

Description:
You can define structs without any fields! These are called unit-like structs.
They are useful when you need a type to implement a trait on, but don't need to store data.

Your task is to define a unit struct `Unit` and instantiate it.
*/

// TODO: Define unit struct Unit


fn main() {
    // TODO: Instantiate it
    // let u = Unit;

    println!("Success!");
}

#[cfg(test)]
mod tests {
    #[test]
    fn test_main_runs() {
        super::main();
    }
}
