// I AM NOT DONE

/*
Difficulty: ⭐⭐
Topic: Structs - Tuple Structs

Description:
Tuple structs are like tuples but have a name.
The fields are accessed by index (0, 1, 2) instead of name.

Your task is to define a tuple struct `Point` with two `i32` fields and create an instance.
*/

// TODO: Define tuple struct Point
// struct Point(...)

fn main() {
    // TODO: Create an instance
    // let p = Point(10, 20);

    // println!("Point: {}, {}", p.0, p.1);
}

#[cfg(test)]
mod tests {
    #[test]
    fn test_main_runs() {
        super::main();
    }
}
