// I AM NOT DONE

/*
Difficulty: ⭐⭐⭐⭐
Topic: Functions

Description:
⭐⭐⭐⭐ - Combine all concepts (build a mini calculator)
*/

fn main() {
    // TODO: Fix this code
    let x = "change me";
    println!("Exercise: {}", x);
}

#[cfg(test)]
mod tests {
    #[test]
    fn test_main_runs() {
        super::main();
    }
}
