// I AM NOT DONE

/*
Difficulty: ⭐
Topic: Functions - Definition

Description:
Functions are the primary way code is executed in Rust.
They are declared using the `fn` keyword.
Rust uses "snake_case" as the conventional style for function names (lowercase with underscores).

The `main` function is special: it's the entry point of every executable Rust program.

Your task is to:
1. Define a new function named `call_me` outside of `main`.
2. Call this function inside `main`.

Hints:
1. Function definition: `fn function_name() { ... }`
2. Function call: `function_name();`
*/

fn main() {
    // TODO: Call the function `call_me`
}

// TODO: Define the function `call_me`
