// I AM NOT DONE

/*
Difficulty: ⭐⭐
Topic: HashSet

Description:
`HashSet` stores unique values. It's like a HashMap with only keys.

Your task is to:
1. Create a `HashSet` called `books`.
2. Insert "Harry Potter" twice.
3. Assert that `len()` is 1.
*/

use std::collections::HashSet;

fn main() {
    // TODO: Create books
    // let mut books = ...;

    // TODO: Insert "Harry Potter" twice

    // println!("Books: {:?}", books);
    // assert_eq!(books.len(), 1);
}

#[cfg(test)]
mod tests {
    #[test]
    fn test_main_runs() {
        super::main();
    }
}
