// I AM NOT DONE

/*
Difficulty: ⭐⭐⭐⭐⭐
Topic: Collections

Description:
⭐⭐⭐⭐⭐ - Build a simple phone book or inventory system
*/

fn main() {
    // TODO: Fix this code
    let x = "change me";
    println!("Exercise: {}", x);
}

#[cfg(test)]
mod tests {
    #[test]
    fn test_main_runs() {
        super::main();
    }
}
