// I AM NOT DONE

/*
Difficulty: ⭐⭐
Topic: Vectors - Iterating

Description:
You can iterate over a vector using a for loop.
If you want to read values, iterate over `&v`.

Your task is to iterate over `v` and print each element.
*/

fn main() {
    let v = vec![100, 32, 57];

    // TODO: Iterate over `v`
    // for i in ... {
    //    println!("{}", i);
    // }
}

#[cfg(test)]
mod tests {
    #[test]
    fn test_main_runs() {
        super::main();
    }
}
