// I AM NOT DONE

/*
Difficulty: ⭐⭐⭐
Topic: Vectors - Storing Different Types

Description:
Vectors can only store one type.
To store multiple types, you can use an Enum wrapper.

Your task is to:
1. Define an enum `SpreadsheetCell` with variants `Int(i32)` and `Float(f64)`.
2. Create a vector `row` containing one Int(3) and one Float(10.12).
*/

// TODO: Define SpreadsheetCell

fn main() {
    // TODO: Create the vector
    // let row = vec![
    //     SpreadsheetCell::Int(3),
    //     SpreadsheetCell::Float(10.12),
    // ];

    // println!("Row: {:?}", row);
}

#[cfg(test)]
mod tests {
    #[test]
    fn test_main_runs() {
        super::main();
    }
}
