// I AM NOT DONE

/*
Difficulty: ⭐
Topic: Testing - Assert Equality

Description:
`assert_eq!` compares two values for equality.
It prints both values if the assertion fails, which makes debugging easier.

Your task is to fix the test.
*/

#[cfg(test)]
mod tests {
    #[test]
    fn you_can_assert_eq() {
        // TODO: Fix the assertion
        assert_eq!(2 + 2, 5);
    }
}
