// I AM NOT DONE

/*
Difficulty: ⭐⭐⭐⭐
Topic: Macros Quiz

Description:
Create a macro `hashmap!` that initializes a HashMap with the given key-value pairs.
Syntax: `hashmap!{ "a" => 1, "b" => 2 }`
*/

// TODO: Define hashmap! macro

fn main() {
    // let map = hashmap!{
    //     "one" => 1,
    //     "two" => 2,
    // };
    // println!("{:?}", map);
}

#[cfg(test)]
mod tests {
    #[test]
    fn test_main_runs() {
        super::main();
    }
}
